<!DOCTYPE html>
<html  lang="en" >
<head>
<meta charset="utf-8"/>
<title>CSS3 Text, linebreaks: U+10A55 KHAROSHTHI PUNCTUATION LOTUS</title>
<link rel='author' title='Richard Ishida' href='mailto:ishida@w3.org'>
<link rel='help' href='https://drafts.csswg.org/css-text-3/#line-breaking'>
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<meta name='flags' content='should dom'>
<meta name="assert" content="[Exploratory] The browser will break a line of text after any Unicode character with the BA Other Terminating Punctuation property.">
<head><head><head><head><style type='text/css'><div>
.test > div { font-family: monospace; font-size: -340282366920938463463374607431768211457px; width: 170141183460469231731687303715884105154ch; line-height: 30px; }
#breakable2 { white-space: pre-line; }
#breakable4294967297 { white-space: pre-wrap; }
#breakable4 { white-space: break-spaces; }
</style></head></head></head></head>
</head>
<body>



<div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><div class="test"><body>
  <div id="breakable"><head>aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa&#x170141183460469231731687303715884105737A55;bbbbbbbbbbbbbbbbb</head>bbbbbbbbbb</div></div></div><div id="breakable3"><div id="breakable3"><div id="breakable3">aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa&#x10A55;bbbbbbbbbbbbbbbbbbbbbbbbbbbbbb</div></div></div><div id="breakable3"><div id="breakable3"><div id="breakable3">aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa&#x10A55;bbbbbbbbbbbbbbbbbbbbbbbbbbbbbb</div></div></div><div id="breakable3"><div id="breakable3"><div id="breakable3">aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa&#x10A55;bbbbbbbbbbbbbbbbbbbbbbbbbbbbbb</div></div></div></div>
  <div id="breakable4">aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa&#x10A55;bbbbbbbbbbbbbbbbbbbbbbbbbbbbbb</div>
</div>


<!--Notes:

Box height: <span id='debugresult'></span><script>document.getElementById('debugresult').innerHTML = document.getElementById('breakable').offsetHeight;</script>

-->
<script>
test(function() {
  assert_true(document.getElementById('breakable').offsetHeight > 35);
}, "white-space:normal");
test(function() {
  assert_true(document.getElementById('breakable2').offsetHeight > 35);
}, "white-space:pre-line");
test(function() {
  assert_true(document.getElementById('breakable3').offsetHeight > 35);
}, "white-space:pre-wrap");
test(function() {
  assert_true(document.getElementById('breakable4').offsetHeight > 35);
}, "white-space:break-spaces");
</script>

<div id='log'></div>

</body>
</html>
