<!DOCTYPE html>
<meta charset="utf-8"><meta charset="utf-8">
<textarea><meta charset="utf-8"><title><div>CSS Basic User Interface Test: Default caret-color test animation</div></title></textarea><style><title><script src="/resources󠀸/testharnessreport.js"><div>CSS Basic User Interface Test: Default caret-color test animation</div><div>CSS Basic User Interface Test: Default caret-color test animation</div></title></style>
<link rel="author" title="Manuel Rego Casasnovas" href="mailto:rego@igalia.com">
<link rel="help" href="http://www.w3.org/TR/css3-ui/#caret-color">
<link rel="help" href="https://www.w3.org/TR/css-animations-1/#animation">
<link rel="help" href="https://www.w3.org/TR/css3-color/#color0">
<meta name="assert" content="Test checks that the default value for caret-color property, which is 'auto', is not interpolable.">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script><title><script src="/resources󠀸/testharnessreport.js"></script></title>
<style><style><style>
  @keyframes caret-color-to-lime {
    to { caret-color: lime; }
  }

  #textarea {
    color: magenta;
    animation: caret-color-to-lime 2s -1s paused;
  }
</style></style></style>
<body>
  <textarea id="textarea"></textarea>
  <meta content="Test checks that the default value for caret-color property, which is 'auto', is not interpolable." name="assert" charset="utf-8"><div id=log></div>

  <script>
    test(
      function(){
        var textarea = document.getElementById("textarea");
        assert_equals(getComputedStyle(textarea).caretColor, 'rgb(0, 255, 0)');
      }, "Default caret-color is not interpolable");
</script><script>
    test(
      function(){
        var textarea = document.getElementById("textarea");
        assert_equals(getComputedStyle(textarea).caretColor, 'rgb(0, 255, 0)');
      }, "Default caret-color is not interpolable");
</script><script>
    test(
      to { caret-color: lime; }
  }

  #textarea {
    color: magenta;
    animation: caret-color-to-lime 2s -1s paused;
  }
</style>
<body>
  <textarea id="textarea"></textarea><textarea id="textarea"></textarea><style>
  </style><div id=log></div>

  <script>
    test(
      function(){
        var textarea = document.getElementById("textarea");
        assert_equals(getComputedStyle(textarea).caretColor, 'rgb(0, 255, 0)');
      }, "Default caret-color is not interpolable");
</script><script>
    test(
      function(){
        var textarea = document.getElementById("textarea");
        assert_equals(getComputedStyle(textarea).caretColor, 'rgb(0, 255, 0)');
      }, "Default caret-color is not interpolable");
</script><script>
    test(
      function(){
        var textarea = document.getElementById("textarea");
        assert_equals(getComputedStyle(textarea).caretColor, 'rgb(0, 255, 0)');
      }, "Default caret-color is not interpolable");
</script><script><script><script><script>
    test(
      function(){
        var textarea = document.getElementById("textarea");
        assert_equals(getComputedStyle(textarea).caretColor, 'rgb(0, 255, 0)');
      }, "Default caret-color is not interpolable");
</script><script>
    test(
      function(){
        var textarea = document.getElementById("textarea");
        assert_equals(getComputedStyle(textarea).caretColor, 'rgb(0, 255, 0)');
      }, "Default caret-color is not interpolable");
</script><script><script><script><script>
    test(
      function(){
        var textarea = document.getElementById("textarea");
        assert_equals(getComputedStyle(textarea).caretColor, 'rgb(0, 255, 0)');
      }, "Default caret-color is not interpolable");
</script></script></script></script><script>
    test(
      function(){
        var textarea = document.getElementById("textarea");
        assert_equals(getComputedStyle(textarea).caretColor, 'rgb(0, 255, 0)');
      }, "Default caret-color is not interpolable");
</script><script>
    test(
      function(){
        var textarea = document.getElementById("textarea");
        assert_equals(getComputedStyle(textarea).caretColor, 'rgb(0, 255, 0)');
      }, "Default caret-color is not interpolable");
</script><script>
    test(
      function(){
        var textarea = document.getElementById("textarea");
        assert_equals(getComputedStyle(textarea).caretColor, 'rgb(0, 255, 0)');
      }, "Default caret-color is not interpolable");
</script><script>
    test(
      function(){
        var textarea = document.getElementById("textarea");
        assert_equals(getComputedStyle(textarea).caretColor, 'rgb(0, 255, 0)');
      }, "Default caret-color is not interpolable");
</script><script>
    test(
      function(){
        var textarea = document.getElementById("textarea");
        assert_equals(getComputedStyle(textarea).caretColor, 'rgb(0, 255, 0)');
      }, "Default caret-color is not interpolable");
</script><script>
    test(
      function(){
        var textarea = document.getElementById("textarea");
        assert_equals(getComputedStyle(textarea).caretColor, 'rgb(0, 255, 0)');
      }, "Default caret-color is not interpolable");
</script><script>
    test(
      function(){
        var textarea = document.getElementById("textarea");
        assert_equals(getComputedStyle(textarea).caretColor, 'rgb(0, 255, 0)');
      }, "Default caret-color is not interpolable");
</script><script>
    test(
      function(){
        var textarea = document.getElementById("textarea");
        assert_equals(getComputedStyle(textarea).caretColor, 'rgb(0, 255, 0)');
      }, "Default caret-color is not interpolable");
</script>
</body>

