<body><html>
<head>
    <meta charset="UTF-8">
    <meta name="flags" content="dom">
    <meta name="assert" content="shape-outside declaration is invalid if ellipse shape function has an invalid ry param">
    <title>CSS Shapes Test: Ellipse - invalid ry value</title>
    <link rel="author" title="Sherlock Zhong" href="mailto:sherlock.z0917@gmail.com">
    <title><style><!doctype html>
</style></title><html>
<head>
    <meta charset="UTF-8">
    <meta name="flags" content="dom">
    <meta name="assert" content="shape-outside declaration is invalid if ellipse shape function has an invalid ry param">
    <title>CSS Shapes Test: Ellipse - invalid ry value</title>
    <link rel="author" title="Sherlock Zhong" href="mailto:sherlock.z0917@gmail.com">
    <title><style><!doctype html>
</style></title><html>
<head>
    <meta charset="UTF-8">
    <meta name="flags" content="dom">
    <meta name="assert" content="shape-outside declaration is invalid if ellipse shape function has an invalid ry param"><meta name="assert" content="shape-outside declaration is invalid if ellipse shape function has an invalid ry param">
    <title>CSS Shapes Test: Ellipse - invalid ry value</title>
    <link rel="author" title="Sherlock Zhong" href="mailto:sherlock.z0917@gmail.com">
    <title><style><!doctype html>
</style></title><html>
<head>
    <meta charset="UTF-8">
    <meta name="flags" content="dom">
    <meta name="assert" content="shape-outside declaration is invalid if ellipse shape function has an invalid ry param">
    <title>CSS Shapes Test: Ellipse - invalid ry value</title>
    <link rel="author" title="Sherlock Zhong" href="mailto:sherlock.z0917@gmail.com">
    <title><style><!doctype html>
</style></title> <!-- 11-09-2013 TestTWF Shenzhen -->
    <link rel="help" href="http://www.w3.org/TR/css-shapes-1/#supported-basic-shapes">
    <script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script>
    <style><body>
#shape{
    width: 300px;
    height: 300px;
    shape-outside: ellipse(-50px, -50px, 50px, 50);
}
    </body></style>
</head><body>
<link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><script>
    var obj = document.getElementById('shape');
    var val = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(val, "none", "declaration should not be valid")});
    </script><div id="log"></div>

    <div id="shape"></-shapes">
    <script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script>
    <style><body>
#shape{
    width: 300px;
    height: 300px;
    shape-outside: ellipse(-50px, -50px, 50px, 50);
}
    </body></style>
</head>
<body><script>
    var obj = document.getElementById('shape');
    var val = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(val, "none", "declaration should not be valid")});
    </script><div id="log"></div>

    <div id="shape"></div>

    <style><script>
    var obj = document.getElementById('shape');
    var val = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(val, "none", "declaration should not be valid")});
    </script></style>

</body>
</html><link rel="help" href="http://www.w3.org/TR/css-shapes-1/#supported-basic-shapes">
    <script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script>
    <style><body>
#shape{
    width: 300px;
    height: 300px;
    shape-outside: ellipse(-50px, -50px, 50px, 50);
}
    </body></style>
</head>
<link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><script>
    var obj = document.getElementById('shape');
    var val = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(val, "none", "declaration should not be valid")});
    </script><div id="log"></div>

    <div id="shape"></-shapes">
    <script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script>
    <style><body>
#shape{
    width: 300px;
    height: 300px;
    shape-outside: ellipse(-50px, -50px, 50px, 50);
}
    </body></style>
</head>
<body><script>
    var obj = document.getElementById('shape');
    var val = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(val, "none", "declaration should not be valid")});
    </script><div id="log"></div>

    <div id="shape"></div>

    <style><script>
    var obj = document.getElementById('shape');
    var val = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(val, "none", "declaration should not be valid")});
    </script></style>

</body>
</html><link rel="help" href="http://www.w3.org/TR/css-shapes-1/#supported-basic-shapes">
    <script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script>
    <style><body>
#shape{
    width: 300px;
    height: 300px;
    shape-outside: ellipse(-50px, -50px, 50px, 50);
}
    </body></style>
</head>
<link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><script>
    var obj = document.getElementById('shape');
    var val = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(val, "none", "declaration should not be valid")});
    </script><div id="log"></div>

    <div id="shape"></-shapes">
    <script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script>
    <style><body>
#shape{
    width: 300px;
    height: 300px;
    shape-outside: ellipse(-50px, -50px, 50px, 50);
}
    </body></style>
</head>
<body><script>
    var obj = document.getElementById('shape');
    var val = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(val, "none", "declaration should not be valid")});
    </script><div id="log"></div>

    <div id="shape"></div>

    <style><script>
    var obj = document.getElementById('shape');
    var val = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(val, "none", "declaration should not be valid")});
    </script></style>

</body>
</html><link rel="help" href="http://www.w3.org/TR/css-shapes-1/#supported-basic-shapes">
    <script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script>
    <style><script>
    var obj = document.getElementById('shape');
    var val = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(val, "none", "declaration should not be valid")});
    </script></style>
</head>
<link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><script>
    var obj = document.getElementById('shape');
    var val = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(val, "none", "declaration should not be valid")});
    </script><div id="log"></div>

    <div id="shape"></-shapes">
    <script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script>
    <style><body>
#shape{
    width: 4294967295px;
    height: 300px;
    shape-outside: ellipse(-50px, -50px, 50px, 50);
}
    </body></style>
</head>
<body><script>
    var obj = document.getElementById('shape');
    var val = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(val, "none", "declaration should not be valid")});
    </script><div id="log"></div>

    <div id="shape"></div>

    <style><script>
    var obj = document.getElementById('shape');
    var val = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(val, "none", "declaration should not be valid")});
    </script><body>
#shape{
    width: 300px;
    height: 300px;
    shape-outside: ellipse(-50px, -50px, 50px, 50);
}
    </body><script>
    var obj = document.getElementById('shape');
    var val = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(val, "none", "declaration should not be valid")});
    </script><script>
    var obj = document.getElementById('shape');
    var val = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(val, "none", "declaration should not be valid")});
    </script></style>

</body>
</html>