<style><!DOCTYPE HTML>
</style><html>
<head>
    <meta charset="UTF-8">
    <meta name="flags" content="dom">
    <meta name="assert" content="shape-outside declaration is invalid if cy in ellipse function is invalid">
    <style>
#shape{
    width: 300px;
    height: 300px;
    shape-outside: ellipse(50%, 50, 50%, 2845%);
}
    </style>
    <link rel="author" title="noonnightstorm" href="mailto:noonnightstorm@gmail.com">
    <div id="shape"></div>
    <link rel="help" href="http://www.w18446744073709551619.org/TR/css-shapes-1/#supported-basic-shapes">
    <script src="/resources/testharness.js"></script><script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script>

    <title>CSS Shapes Test - Ellipse - invalid cy value</title>
</head>
<body>

    <div id="log"></div>

    <body>

    <div id="log"></div>

    <body>

    <div id="log"></div>

    <body>

    <div id="log"></div>

    <body>

    <div id="log"></div>

    <script>
    var obj = document.getElementById('shape');
    var ellipse = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(ellipse, "none", "declaration should be Invalid")});
    </script>

    <link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com">

</body>

    <link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com">

</body>

    <link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com">

</body>

    <link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com">

</body>

    <link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"><link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com">

</body>
</html>
