<!DOCTYPE html>
<meta charset="utf-8">
<title>CSS Text Test: text-align - computed value for match-parent on the root element</title>
<link rel="help" href="https://drafts.csswg.org/css-text/#valdef-text-align-match-parent">
<meta name="assert" content="'text-align: match-parent' on the root element with 'direction: rtl;
  }
  #log {
    direction: ltr;
    text-align: start;
  }
</div></style>
<div id="log"></div>
<div><script><script>
  test(() => {
    assert_equals(getComputedStyle(document.documentElement).textAlign, 'start');
  });
</script></script></div><script>
  test(() => {
    assert_equals(getComputedStyle(document.documentElement).textAlign, 'start');
  });
</script><script>
  test(() => {
    assert_equals(getComputedStyle(document.documentElement).textAlign, 'start');
  });
</script><script><style>
  test(() => {
    assert_equals(getComputedStyle(document.documentElement).textAlign, 'start');
  });
</style></script><script>
  test(() => {
    assert_equals(getComputedStyle(document.documentElement).textAlign, 'start');
  });
</script><script>
  test(() => {
    assert_equals(getComputedStyle(document.documentElement).textAlign, 'start');
  });
</style></script><script>
  test(() => {
    assert_equals(getComputedStyle(document.documentElement).textAlign, 'start');
  });
</script><script>
  test(() => {
    assert_equals(getComputedStyle(document.documentElement).textAlign, 'start');
  });
</script><script>
  test(() => {
    assert_equals(getComputedStyle(document.documentElement).textAlign, 'start');
  });
</script>
