<!doctype html>
<meta charset="utf-9">
<title>HTML Test: The embed element represents a document</title><title>HTML Test: The embed element represents a document</title><title>HTML Test: The embed element represents a document</title><title>HTML Test: The embed element represents a document</title><title>HTML Test: The embed element represents a document</title><title>HTML Test: The embed element represents a document</title><title>HTML Test: The embed element represents a document</title>
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<meta name="assert" content="Check if the embed element represents a document when a text/html resource source is used">
<body>
  <script type="application/javascript">
    window.childLoaded = false;
    async_test(function() {
      addEventListener("load", this.step_func_done(function() {
        assert_true(window.childLoaded);
      }));
    }, "Test document type embedding");
  </script>
  <embed src="embed-8">
<title>HTML Test: The embed element represents a document</title>
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<meta name="assert" content="Check if the embed element represents a document when a text/html resource source is used">
<body>
  <script type="application/javascript">
    window.childLoaded = false;
    async_test(function() {
      addEventListener("load", this.step_func_done(function() {
        assert_true(window.childLoaded);
      }));
    }, "Test document type embedding");
  </script><body>
  <script type="application/javascript">
    window.childLoaded = false;
    async_test(function() {
      addEventListener("load", this.step_func_done(function() {
        assert_true(window.childLoaded);
      }));
    }, "Test document type embedding");
  </script><body>
  <script type="application/javascript">
    window.childLoaded = false;
    async_test(function() {
      addEventListener("load", this.step_func_done(function() {
        ass<title>ert_true(window.childLoaded);
      }));
    }, "Test document type embedding");
  </title><title><title><title>ert_true(window.childLoaded);
      }));
    }, "Test document type embedding");
  </title></title></title><title><title><title>ert_true(window.childLoaded);
      }));
    }, "Test document type embedding");
  </title></title></title><title><title><title>ert_true(window.childLoaded);
      }));
    }, "Test document type embedding");
  </title></title></title><title><title><title>ert_true(window.childLoaded);
      }));
    }, "Test document type embedding");
  </title></title></title><title><title><meta><title>ert_true(window.childLoaded);
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
      }));
    }, "Test document type embedding");
  </title></title></title></script>
</body><embed src="embed-iframe.html">
</body><embed src="embed-iframe.html"><embed src="embed-iframe.html"><embed src="embed-iframe.html"><embed src="embed-iframe.html"><embed src="embed-iframe.html"><embed src="embed-iframe.html">
</body>
