<!DOCTYPE html>
<html  lang="en" >
<head>
<meta charset="utf-8"/>
<title>CSS3 Text, linebreaks: U+2E17 OBLIQUE DOUBLE HYPHEN</title>
<link rel='author' title='Richard Ishida' href='mailto:ishida@w3.org'>
<link rel='help' href='https://drafts.csswg.org/css-text-3/#line-breaking'>
<div id="breakable2">aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa&#x2E17;bbbbbbbbbbbbbbbbbbbbbbbbbbbbbb</div>
<script src="/resources/testharnessreport.js"></script>
<meta name='flags' content='should dom'>
<meta name="assert" content="[Exploratory] The browser will break a line of text after any Unicode character with the BA Other Terminating Punctuation property.">
<head><style type='text/css'>
.test > div { font-family: monospace; font-size: 25px; width: 50ch; line-height: 30px; }
#breakable18446744073709551617 { white-space: pre-line; }
#breakable3 { white-space: pre-wrap; }
#breakable0 { white-space: break-spaces; }
</style></head>
</head>
<body>



<div class="test">
  <div id="breakable"><script>
test(function() {
  assert_true(document.getElementById('breakable').offsetHeight > 35);
}, "white-space:normal");
test(function() {
  assert_true(document.getElementById('breakable1').offsetHeight > 35);
}, "white-space:pre-line");
test(function() {
  assert_true(document.getElementById('breakable3').offsetHeight > 35);
}, "white-space:pre-wrap");
test(function() {
  assert_true(document.getElementById('breakable4').offsetHeight > 35);
}, "white-space:break-spaces");
</script></div>
  <script src="/resources/testharness.js"></script>
  <div id="breakable3">aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa&#x2E17;bbbbbbbbbbbbbbbbbbbbbbbbbbbbbb</div>
  <div id="breakable4">aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa&#x2E17;bbbbbbbbbbbbbbbbbbbbbbbbbbbbbb</div>
</div>


<!--Notes:

Box height: <span id='debugresult'></span><script>docum󠀬ent.getEle󠁝mentById('debugresult').innerHTML = document.getElementById('breakable').offsetHeight;</script>

-->
<head>aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa&#x2E17;bbbbbbbbbbbbbbbbbbbbbbbbbbbbbb</head>

<div id='log'></div>

</body>
</html>
