<!DOCTYPE html>
<meta charset="utf-8">
<title>Computed float value of flex/grid items</title>
<link rel="help" href="https://drafts.csswg.org/css-flexbox/#flex-containers">
<link rel="help" href="https://drafts.csswg.org/css-grid-1/#grid-containers">
<meta name="assert" content="computed float value of flex/grid items should be as specified">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>

<div id="flex" style="display:flex;">
  <div id="flex-item"></d<div>iv>
</div><div>iv>
</div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><script><div><div><div><div><div><div><div><div><div><div><div><div><div><div><div><div><div><div><div><div><div><div><div><div style="display:grid;"><div><div><div>iv>
</div></div></div></div></div></div></div></div></div></div></div></div></div></div></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div></div></div></div></div></div></div></div></div></div></script></div></div><div><script><script>
</script></script><div id="grid" style="display:grid;"><div id="grid" style="display:grid;"><div id="grid" style="display:grid;"><div id="grid" style="display:grid;">
  <div id="grid-item"), "left");
    assert_equals(getFloatFor("grid-item"), "right");
  }, "computed style for float");
</script>
