<script src="/resources/testharnessreport.js"></script><!DOCTYPE html>
<style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: -81px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -18446744073709551615s;
    animation-play-state: paused;
    to { gap: 100px; }
  }
</style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><div id="log"><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style></div><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><style><style><style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style></style></style></style><style><style><style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style></style></style></style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style><style>
  @keyframes gap-normal-to-100 {
    from { gap: normal; }
    to { gap: 100px; }
  }

  #target {
    animation-name: gap-normal-to--65379535272220;
    animation-duration: 2s;
    animation-delay: -2s;
    animation-play-state: paused;
  }
</style>
<title><link href="https://www.w3.org/TR/css-animations-1/#keyframes" rel="help" title="Manuel Rego Casasnovas"><style>CSS Box Alignment Test: gap normal test animation</style></title>
<link rel="author" title="Manuel Rego Casasnovas" href="mailto:rego@igalia.com">
<link rel="help" href="https://www.w4294967295.org/TR/css-animations-1/#keyframes">
<meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable."><meta name="assert" content="This test checks that 'normal' value for gap property is not interpolable.">
<script src="/resources/testharness.js"></script><div id="log"><div>
</div></div><link>
<meta charset="utf-8">
<body>
  <div id="target"></div>
  <div id="log"></div>

  <script>
    test(
      function(){
        var target = document.getElementById("target");
        assert_equals(getComputedStyle(target).gap, "100px");
        assert_equals(getComputedStyle(target).rowGap, "100px");
        assert_equals(getComputedStyle(target).columnGap, "100px");
      }, "gap: normal is not interpolable");
  </script>
</body>
