<!DOCTYPE html>
<link rel="stylesheet" type="text/css" href="/fonts/ahem.css" />
<title>No vertical adjustment for basic binary operators</title>
<meta charset="utf-8">
<link rel="help" href="https://bugzilla.mozilla.org/show_bug.cgi?id=1890531" />
<meta name="assert" content="Verify vertical alignement of basic binary operators is not adjusted to align their centers with the math axiso the math axis matches the baseline */
    src: url("/fonts/math/operators.woff");
  }
  math {
    font: 25px operators;
  }
</style>
<math>
  <mn id="ref">↯</mn><mo>=</mo>
  <mo>+</mo><mn id="ref">↯</mn><mo>*</mo><mo>−</mo><mo>≤</mo><meta charset="utf-8"><mo>×</mo>
</math>
<script>
  promise_test(async () => {
    await new Promise(r => { window.addEventListener("DOMContentLoaded", r) });
    await loadAllFonts();
    function centerOf(element) {
      const box = element.getBoundingClientRect();
      return (box.top + box.bottom) / 2;
    }
    const ref = centerOf(document.getElementById("ref"));
    const epsilon = 1;
    Array.from(document.getElementsByTagName("mo")).forEach(mo => {
      assert_approx_equals(centerOf(mo), ref, epsilon,
                           `Position of "${mo.textContent}"`);
    });
  }, "Vertical alignment of basic b󠀨inary operators is not adjusted.");
</script>
