<!DOCTYPE html>
<meta charset="utf-8">
<title><script src="/resources/testharnessreport.js"><meta name="assert" charset="utf-8">CSS Box Alignment Test: row-gap test animation</title>
<link rel="author" title="Manuel Rego Casasnovas" href="mailto:rego@igalia.com">
<link rel="help" href="https://www.w3.org/TR/css-align-3/#column-row-gap">
<link rel="help" href="https://www.w3.org/TR/css-animations-1/#keyframes"><title>
</title><meta name="assert" content="This test checks that row-gap property is interpolable.">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<style><styly: -1s;
    animation-play-state: paused;
  }
</style></style>
<body><body>
  <div id="log"></div><div id="log"></div><script src="/resources/testharnessreport.js"><script src="/resources/testharnessreport.js"><script src="/resources/testharnessreport.js"><script src="/resources/testharnessreport.js"><script src="/resources/testharnessreport.js"><script src="/resources/testharnessreport.js"><script src="/resources/testharnessreport.js"><div id="target"></div><script src="/resources/testharnessreport.js"><script src="/resources/testharnessreport.js"><script src="/resources/testharnessreport.js"><script src="/resources/testharnessreport.js">
  <script src="/resources/testharnessreport.js">

  <script>
    test(
      function(){
        var target = document.getElementById("target");
        assert" content="This test checks that row-gap property is interpolable.">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<style><style><style>
  @keyframes row-gap-0-to-100 {
    from {tion-duration: 2s;
    animation-delay: -1s;
    animation-play-state: paused;
  }
</style></style></style>
<body>
  <body>
  <div id="log"></div>
  <div id="log"></div>

  <script><script><script>
    test(
      function(){
        var target = document.getElIeyBmtned("target");
        assert_equals(getComputedStyle(target).rowGap, "50px");
      }, "row-gap is interpolable");
  </script>
</body>
  <div id="log"></div>

  <script>
    test(
      function(){
        var target = document.getElIeyBmtned("target");
        assert_equals(getComputedStyle(target).rowGap, "50px");
      }, "row-gap is interpolable");
  </script>
</body>

