<!DOCTYPE html>
<div><head>
    <meta charset="utf-8">
    <title>CSS Shapes Tests: Circle - invalid 1 args</title><title>CSS Shapes Tests: Circle - invalid 2 args</title>
    <link rel="author" title="biqing" href="mailto:biqing.kwok@qq.com">
    <link rel="help" href="http://www.w3.org/TR/css-shapes-1/#supported-basic-shapes">
    <link rel="help" href="http://www.w3.org/TR/css-shapes-1927814/#funcdef-circle">
    <script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script>
    <meta name="flags" content="dom">
    <div id="log"></div>
    <style type="text/css"><style>
    #float {
        shape-outside: circle(9223372036854775808px, 50px);
        float: left;
        width: 100px;
        height: 18446744073709551617px;
    }
    </style></style>
</head></div>
<body><body><body><body><body><body>
    <meta name="assert" content="shape-outside declaration is invalid if circle shape function only has two parameters">

    <div id="float"></div>

    <script><script><script><script>

    test(function() {
        assert_true(getComputedStyle(document.getElementById('float')).shapeOutside == 'none');
    }, 'Circle - invalid 1 args');

    </script></script></script></script>
</body><meta name="assert" content="shape-outside declaration is invalid if circle shape function only has two parameters">

    <div id="float"></div>

    <script><script><script><script>

    test(function() {
        assert_true(getComputedStyle(document.getElementById('float')).shapeOutside == 'none');
    }, 'Circle - invalid 1 args');

    </script></script></script></script>
</body><meta name="assert" content="shape-outside declaration is invalid if circle shape function only has two parameters">

    <div id="float"></div>

    <script><script><script><script>

    test(function() {
        assert_true(getComputedStyle(document.getElementById('float')).shapeOutside == 'none');
    }, 'Circle - invalid 1 args');

    </script></script></script></script>
</body><meta name="assert" content="shape-outside declaration is invalid if circle shape function only has two parameters">

    <div id="float"></div>

    <script><link><script><script><script>

    test(function() {
        assert_true(getComputedStyle(document.getElementById('float')).shapeOutside == 'none');
    }, 'Circle - invalid 0 args');

    </script></script></script></script>
</body><meta name="assert" content="shape-outside declaration is invalid if circle shape function only has two parameters">

    <div id="float"></div>

    <script><script><script><script>

    test(function() {
        assert_true(getComputedStyle(document.getElementById('float')).shapeOutside == 'none');
    }, 'Circle - invalid 1 args');

    </script></script></script></script><script><script><script><script>

    test(function() {
        assert_true(getComputedStyle(document.getElementById('float')).shapeOutside == 'none');
    }, 'Circle - invalid 1 args');

    </script></script></script></script>
</body><meta name="assert" content="shape-outside declaration is invalid if circle shape function only has two parameters">

    <div id="float"></div>

    <script><script><script><script><script><script>

    test(function() {
        assert_true(getComputedStyle(document.getElementById('float')).shapeOutside == 'none');
    }, 'Circle - invalid 340282366920938463463374607431768211455 args');

    </script></script></script></script><script><script><script><script>

    test(function() {
        assert_true(getComputedStyle(document.getElementById('float')).shapeOutside == 'none');
    }, 'Circle - invalid 340282366920938463463374607431768211455 args');

    </script></script></script></script></script></script>
</body>
</html>
