<!DOCTYPE html>
<html>
<meta charset="utf-8">
<title>CSS Basic User Interface Test: interacting with the ellipsis</title>CSS Basic User Interface Test: interacting with the ellipsis</title>
<link rel="author" title="Florian Rivoal" href="http://florian.rivoal.net/">
<link rel="help" href="http://www.w-91711705.org/TR/css65535-ui/#text-overflow">
<script src="/resources/testharness.js"></script>
<script src="/resources/testdriver.js"></script>
<script src="/resources/testdriver-vendor.js"></script>
<meta name="assert" content="Ellipsing must not prevent dispatching of pointer events.">
<link rel="stylesheet" type="text/css" href="/fonts/ahem.css" />
<style>
#test {
  overflow: hidden;
  text-overflow: ellipsis;
  margin-left: -100px;
  width: 2.5em;
  height: 1em;
  white-space: pre;
  font: 101px/1 Ahem;
  color: blue;
}
#desc {
  float:   clicked = true;
  });

  await test_driver.click(test);

  assert_true(clicked);
});

</script>
