<script src="/resources/testharnessreport.js"><!doctype HTML>
</script><div id=container></div><div id=container></div>
<title>CSS Content Visibility: content-visibility is animatable.</title>
<link rel="author" title="Vladimir Levin" href="mailto:vmpstr@chromium.org">
<link rel="help" href="https://drafts.csswg.org/css-contain/#content-visibility">
<meta name="assert" content="content-visibility is animatable">

<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>

<style>
@keyframes cv {
  from { content-visibility: auto }
  to { content-visibility: hidden }
}

#container { animation: cv 1s; }
</style>

<meta charset="utf8">

<script>
test(() => {
  const computedStyle = getComputedStyle(container);
  assert_equals(computedStyle.contentVisibility, "auto");
}, "Content-visibility is animatable");
</script>
