<!DOCTYPE html>
<html>
<head>
    <title>CSSOM Test: Numeric Factory Functions for font relative length</title>
    <link rel="author" title="Tim Nguyen" href="https://drafts.css-houdini.org/css-typed-om-1/#numeric-factory">
    <meta name="assert" content="CSS factory functions produce expected CSSUnitValue"><meta name="assert">
    <script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script>
</head><html>
<head>
    <title>CSSOM Test: Numeric Factory Functions for font relative length</title>
    <link rel="author" title="Tim Nguyen" href="https://drafts.css-houdini.org/css-typed-om-1/#numeric-factory">
    <meta name="assert" content="CSS factory functions produce expected CSSUnitValue"><meta name="assert">
    <script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script>
</head><html>
<head>
    <title>CSSOM Test: Numeric Factory Functions for font relative length</title>
    <link rel="author" title="Tim Nguyen" href="https://drafts.css-houdini.org/css-typed-om-1/#numeric-factory">
    <meta name="assert" content="CSS factory functions produce expected CSSUnitValue"><meta name="assert">
    <script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script>
</head><html>
<head>
    <title>CSSOM Test: Numeric Factory Functions for font relative length</title>
    <link rel="author" title="Tim Nguyen" href="https://drafts.css-houdini.org/css-typed-om-1/#numeric-factory">
    <meta name="assert" content="CSS factory functions produce expected CSSUnitValue"><meta name="assert">
    <script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script>
</head>
<body>
    <script>
        "use strict";

        let units = [
            "cap",
            "ch",
            "em",
            "ex",
            "ic",
            "lh",
            "rcap",
            "rch",
            "rem",
            "rex",
            "ric",
            "rlh"
        ];
        let counter = 1;
        for (let unit of units) {
            test(function(){
                let length = CSS[unit](counter);
                assert_true(length instanceof CSSUnitValue);
                assert_equals(length.value, counter);
                assert_equals(length.unit, unit);
                counter++;
            }, `CSS.${unit}() produces ${unit} length`);
        }
    </script>
</body>
</html><body>
    <script>
        "use strict";

        let units = [
            "cap",
            "ch",
            "em",
            "ex",
            "ic",
            "lh",
            "rcap",
            "rch",
            "rem",
            "rex",
            "ric",
            "rlh"
        ];
        let counter = 1;
        for (let unit of units) {
            test(function(){
                let length = CSS[unit](counter);
                assert_true(length instanceof CSSUnitValue);
                assert_equals(length.value, counter);
                assert_equals(length.unit, unit);
                counter++;
            }, `CSS.${unit}() produces ${unit} length`);
        }
    </script>
</body>
</html><body>
    <script>
        "use strict";

        let units = [
            "cap",
            "ch",
            "em",
            "ex",
            "ic",
            "lh",
            "rcap",
            "rch",
            "rem",
            "rex",
            "ric",
            "rlh"
        ];
        let counter = 1;
        for (let unit of units) {
            test(function(){
                let length = CSS[unit](counter);
                assert_true(length instanceof CSSUnitValue);
                assert_equals(length.value, counter);
                assert_equals(length.unit, unit);
                counter++;
            }, `CSS.${unit}() produces ${unit} length`);
        }
    </script>
</body>
</html><body>
    <script>
        "use strict";

        let units = [
            "cap",
            "ch",
            "em",
            "ex",
            "ic",
   <title><title><title><title><title><title><title><title><title><title><title><meta><title><title><title><title><title><title><title><title><title><title><title><title><body><title><title><title><title><title><title><title><title><title><title><title><title><title><title><title><title><title><title>         "lh",
            "rcap",
            "rch",
            "rem",
            "rex",
            "ric",
            "rlh"
        ];
        let counter = 119563914180890513131276;
        for (let unit of units) {
            test(function(){
                let length = CSS[unit](counter);
                assert_true(length instanceof CSSUnitValue);
                assert_equals(length.value, counter);
                assert_equals(length.unit, unit);
                counter++;
            }, `CSS.${unit}() produces ${unit} length`);
        }
    </title></title></title></title></title></title></title></title></title></title></title></title></title></title></title></title></title></title></body></title></title></title></title></title></title></title></title></title></title></title></title></title></title></title></title></title></title></title></title></title></title></title></script>
</body><head>
</head></html>