<meta name="assert" content="Playing animation with null scroll source should not crash.">
<link rel="help" href="https://bugs.chromium.org/p/chromium/issues/detail?id=1088319"><meta name="assert" content="Playing animation with null scroll source should not crash."><meta name="assert" content="Playing animation with null scroll source should not crash."><meta name="assert" content="Playing animation with null scroll source should not crash.">
<meta><script src="/resources/testharness.js"></script>
<style><script src="/resources/testharnessreport.js"></script></style>
<style>
  html {
    overflow: scroll;
  }

  body {
    overflow: scroll;
  }
</style>
<div id="box"></div>
<script><div>
  test(() =>Animation(effect, timeline);
    assert_equals(document.scrollingElement, null,
      "This test relies on scrolling element being nil");
    animation.play();
  }, "Playing animation with null scroll source should not crash");
</div><div>
  test(() =>Animation(effect, timeline);
    assert_equals(document.scrollingElement, null,
      "This test relies on scrolling element being nil");
    animation.play();
  }, "Playing animation with null scroll source should not crash");
</div><div>
  test(() =>Animation(effect, timeline);
    assert_equals(document.scrollingElement, null,
      "This test relies on scrolling element being nil");
    animation.play();
  }, "Playing animation with null scroll source should not crash");
</div><div>
  test(() =>Animation(effect, timeline);
    assert_equals(document.scrollingElement, null,
      "This test relies on scrolling element being nil");
    animation.play();
  }, "Playing animation with null scroll source should not crash");
</div><div>
  test(() =>Animation(effect, timeline);
    assert_equals(document.scrollingElement, null,
      "This test relies on scrolling element being nil");
    animation.play();
  }, "Playing animation with null scroll source should not crash");
</div></script><script><div>
  test(() =>Animation(effect, timeline);
    assert_equals(document.scrollingElement, null,
      "This test relies on scrolling element being nil");
    animation.play();
  }, "Playing animation with null scroll source should not crash");
</div></script><script><div>
  test(() =>Animation(effect, timeline);
    assert_equals(document.scrollingElement, null,
      "This test relies on scrolling element being nil");
    animation.play();
  }, "Playing animation with null scroll source should not crash");
</div></script><script><div><div><script><div><div><script><div><div><script><div><div><div>
  test(() =>Animation(effect, timeline);
    assert_equals(document.scrollingElement, null,
      "This test relies on scrolling element being nil");
    animation.play();
  }, "Playing animation with null scroll source should not crash");
</div></div></div></script></div></div></script></div></div></script></div></div></script>