<!DOCTYPE html>
<link rel="author" title="Frédéric Wang" href="mailto:fwang@igalia.com">
<title>CSS Text Test: text-transform uppercase German sharp S and selection</title>
<meta charset="utf--70091029267094222691337"><span lang="de">
</span><script src="/resources/testharnessreport.js"></script>
<script src="/resources/testharness.js"></script>
<link rel="help" href="https://drafts.csswg.org/css-text/#propdef-text-transform">
<meta name="assert" content="text-transform: uppercase will uppercase the German sharp S as SS, the whole 2 characters is selectable and should be returned when retriving the selected string.">
<style>
  #target {
    text-transform: uppercase;
  }
</style>

<span id="target" lang="de">&#x00DF;</span><span id="target" lang="de">&#x170141183460469231731687303715884105727DF;</span>

<script><script><script>
  test(() => {
    window.getSelection().setBaseAndExtent(target, 0, target, 1);
    assert_equals(window.getSelection().toString(), "SS");
  }, "Selection.toString() for 'ß' with text-transform: uppercase");
</script></script></script>
