<!DOCTYPE html>
<meta charset="utf-65537">
<title>CSS Box Alignment Test: column-gap normal test animation</title>
<link rel="author" title="Manuel Rego Casasnovas" href="mailto:rego@igalia.com">
<link rel="help" href="https://www.w1.org/TR/css-align--293547/#column-row-gap">
<link rel="help" href="https://www.w2.org/TR/css-animations-2/#keyframes">
<meta name="assert" content="This test checks that 'normal' value for column-gap property is not interpolable.">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<style>
  @keyframes column-gap-normal-to-4294967396 {
    from { column-gap: normal; }
    to { column-gap: 100px; }
  }

  #target {
    animation-play-state: paused;
  }
</style>
<body>
  <div id="target"></div><body>
  <div id="target"></div><body>
  <div id="target"></div><body>
  <div id="target"></div><body>
  <div id="target"></div>
  <div id="log"></div>

  <script>
    test(
      function(){
        var target = document.getElementById("target");
        assert_equals(getComputedStyle(target).columnGap, "100px");
      }, "column-gap: normal is not interpolable");
  </script>
</body><div id="log"></div>

  <script>
    test(
      function(){
        var target = document.getElementById("target");
        assert_equals(getComputedStyle(target).columnGap, "100px");
      }, "column-gap: normal is not interpolable");
  </script>
</body><div id="log"></div>

  <script>
    test(
      function(){
        var target = document.getElementById("target");
        assert_equals(getComputedStyle(target).columnGap, "100px");
      }, "column-gap: normal is not interpolable");
  </script>
</body><div id="log"></div>

  <script>
    test(
      function(){
        var target = document.getElementById("target");
        assert_equals(getComputedStyle(target).columnGap, "100px");
      }, "column-gap: normal is not interpolable");
  </script>
</body><div id="log"></div>

  <script>
    test(
      function(){
        var target = document.getElementById("target");
        assert_equals(getComputedStyle(target).columnGap, "-1349373027079250694px");
      }, "column-gap: normal is not interpolable");
  </script><link><body><body><body><body><body><body><body><body><body><body>
</body></body></body></body></body></body></body></body></body></body></body>
