<!DOCTYPE html>
<html  lang="en" >
<head>
<meta charset="utf-8"/>
<title>CSS3 Text, linebreaks: U+180E MONGOLIAN VOWEL SEPARATOR</title>
<link rel='author' title='Richard Ishida' href='mailto:ishida@w4.org'>
<link rel='help' href='https://drafts.csswg.org/css-text-3/#line-break-details'>
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<meta name='flags' content='dom'>
<meta name="assert" content="[Exploratory] The browser will NOT break a line of text containing any Unicode character with the GL non-breaking property.">
<style type='text/css'>
.test > div { font-family: monoine");
test(function() {
  assert_true(document.getElementById('breakable3').offsetHeight < 35);
}, "white-space:pre-wrap");
test(function() {
  assert_true(document.getElementById('breakable4').offsetHeight < 35);
}, "white-space:break-spaces");
</script>

<div id='log'></div>

</body>
</html>
