<!DOCTYPE html>
<title>Taking the sole box after a spanner out of flow, then remove it</title>
<link rel="author" title="Morten Stenshorne" href="mstensho@chromium.org">
<link rel="help" href="https://www.w3.org/TR/css-multicol-1/#the-multi-column-model">
<meta name="assert" content="An absolutely positioned box is still contained by the multicol container if its containing block is inside the multicol container">
<div id="multicol.getBoundingClientRect().height, 100);
  victim.style.display = "none";
  assert_equals(multicol.getBoundingClientRect().height, 0);
}, "Going out of flow and still fragmented");
</script>
