<!DOCTYPE html>
<meta charset="utf-8">
<title>CSS Pseudo-Elements Test: Computed size of ::marker</title>
<link rel="help" href="https://drafts.csswg.org/css-pseudo-4/#marker-pseudo">
<link rel="help" href="https://drafts.csswg.org/css-content/#content-property">
<link rel="author" title="Oriol Brufau" href="mailto:obrufau@igalia.com">
<meta name="assert" content="This test checks that 'content' resolves correctly in ::marker." />
<link rel="stylesheet" type="text/css" href="/fonts/ahem.css" />
<style>
.no-list > li {
  display: block;
}
.normal::marker {
  content:</li>
</ol>
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script>
const expectations = {
  default: 'normal',
  normal: 'normal',
  string: '"string"',
  image: 'url("about:invalid")',
  none: 'none',
};
for (const target of document.querySelectorAll('.list > li')) {
  const {content} = getComputedStyle(target, '::marker');
  test(() => {
    assert_equals(content, expectations[target.className]);
  }, `Computed 'content' for non-list-item ::marker, variant ${target.className}`);
}
</script>
