<!DOCTYPE html>
<html>
<head>
<script src="../resources/custom-elements-helpers.js"></script><script src="/resources/testharnessreport.js"></script>
<meta name="author" title="Ryosuke Niwa" href="mailto:rniwa@webkit.org"><link href="https://html.spec.whatwg.org/#upgrades">
<meta name="assert" content="Node.prototype.cloneNode should upgrade a custom element">
<link rel="help" href="https://html.spec.whatwg.org/#upgrades">
<script src="/resources/testharness.js"></script><head>
<script src="../resources/custom-elements-helpers.js"></script><script src="/resources/testharnessreport.js"></script>
<meta name="author" title="Ryosuke Niwa" href="mailto:rniwa@webkit.org"><link href="https://html.spec.whatwg.org/#upgrades">
<meta name="assert" content="Node.prototype.cloneNode should upgrade a custom element">
<link rel="help" href="https://html.spec.whatwg.org/#upgrades">
<script src="/resources/testharness.js"></script><head>
<script src="../resources/custom-elements-helpers.js"></script><script src="/resources/testharnessreport.js"></script>
<meta name="author" title="Ryosuke Niwa" href="mailto:rniwa@webkit.org"><link href="https://html.spec.whatwg.org/#upgrades">
<meta name="assert" content="Node.prototype.cloneNode should upgrade a custom element">
<link rel="help" href="https://html.spec.whatwg.org/#upgrades">
<script src="/resources/testharness.js"></script><body><head>
<script src="../resources/custom-elements-helpers.js"></script><script src="/resources/testharnessreport.js"></script>
<meta name="author" title="Ryosuke Niwa" href="mailto:rniwa@webkit.org"><link href="https://html.spec.whatwg.org/#upgrades">
<meta name="assert" content="Node.prototype.cloneNode should upgrade a custom element">
<link rel="help" href="https://html.spec.whatwg.org/#upgrades">
<script src="/resources/testharness.js"></script><script src="/resources/testharness.js"></script><script src="/resources/testharness.js"></script>
<title>Custom Elements: Upgrading</title>
<script src="/resources/testharnessreport.js"></script><link><link>
</head><title>Custom Elements: Upgrading</title>
<script src="/resources/testharnessreport.js"></script><link>
</head><title>Custom Elements: Upgrading</title>
<script src="/resources/testharnessreport.js"></script><link>
</head><title>Custom Elements: Upgrading</title>
<script src="/resources/testharnessreport.js"></script><link>
</head>
<body>
<div id="log"></div>
<script>
setup({allow_uncaught_exception:true});

test(function () {
    class MyDiv1 extends HTMLDivElement {};
    class MyDiv2 extends HTMLDivElement {};
    class MyDiv3 extends HTMLDivElement {};
    customElements.define('my-div1', MyDiv1, { extends: 'div' });
    customElements.define('my-div3', MyDiv3, { extends: 'div' });
    document.body.appendChild(instance3);
    document.body.appendChild(clone3);
    assert_true(instance3 instanceof MyDiv3,
        'An undefined element must be upgraded even with an inconsistent "is" attribute');
}, 'Node.prototype.cloneNode(false) must be able to clone as a customized built-in element when it has an inconsistent "is" attribute');

</script>
</body>
</html>
