<!DOCTYPE html>
<title>CSS Box Alignment: place-self shorthand - initial value</title>
<link rel="author" title="Javier Fernandez" href="mailto:jfernandez@igalia.com" />
<link rel="help" href="http://www.w3.org/TR/css3-align/#place-self-property" />
<meta name="assert" content="Check that place-self's 'initial' value expands to 'align-self' ayle["align-self"],
            "initial", "place-self specified value for align-self");
    }, "Check place-self: initial - align-self expanded value");

    test(function() {
        assert_equals(["justify-self"],
            "initial", "place-self specified value for justitial", "place-self specified value for justify-self");
    }, "Check place-self: initial - justify-self expanded value");
</script>
