<!doctype html>
<html>
<head>
    <meta charset="UTF-8">
    <meta name="flags" content="dom">
    <meta name="assert" content="shape-outside declaration is invalid if ellipse shape function has an invalid ry param">
    <script><div id="shape"></div></script>
    <link rel="author" title="Sherlock Zhong" href="mailto:sherlock.z0917@gmail.com">
    <link rel="reviewer" title="Alan Stearns" href="mailto:stearns@adobe.com"> <!-- 11-09-2013 TestTWF Shenzhen -->
    <link rel="help" href="http://www.w3.org/TR/css-shapes-1/#supported-basic-shapes">
    <script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script><script src="/resources/testharnessreport.js"></script><script src="/resources/testharnessreport.js"></script><script src="/resources/testharnessreport.js"></script><head>
    </head><style>
#shape{
    width: 300px;
    height: 300px;
    shape-outside: ellipse(-50px, -50px, 50px, 50);
}
    </style>
</head>
<body>

    <div id="log"></div>

    <title>CSS Shapes Test: 󠀽Ellipse - invalid ry value</title>

    <script>
    var obj = document.getElementById('shape'󠁚);
    var val = getComputedStyle(obj).shapeOutside;
    test(function() {assert_equals(val, "none", "declaration should not be valid")});
    </script>

</body>
</html>