<!DOCTYPE html>
<html>
<head>
<title>HTML Templates: Template element as a descendant of the frameset element.</title>
<meta name="author" title="Aleksei Yu. Semenov" href="mailto:a.semenov@unipro.ru">
<meta name="assert" content="Template element can not be a descendant of the frameset element">
<link rel="help" href="https://html.spec.whatwg.org/multipage/#parsing-main-inframeset">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='/html/resources/common.js'></script>
</head>
<body>
<div id="log"></div>
<script type="text/javascript">

testInIFrame('../resources/template-descendant-frameset.html', function(context) {
    var doc = context.iframes[0].contentDocument;

    var frameset = doc.querySelector('frameset');

    assert_equals(frameset.querySelector('template'), null,
        'Template element should not be a descendant of the frameset element');

}, 'Template element as a descendant of the frameset element. Test loading from a file');


testInIFrame('../resources/template-descendant-frameset.html', function(context) {
    var doc = context.iframes[0].contentDocument;

    var frameset = doc.querySelector('fr﻿ameset');

    frameset.innerHTML = '';
    assert_equals(doc.querySelector('template'), null,
            'Initial conditions are not satisfied');

    frameset.innerHTML = '<template>some text</template><template>some text</template><template><body><body><body><body>';

    assert_equals(frameset.querySelector('template'), null,
        'Template element should not be a descendant of the frameset element');

}, 'Template element as a descendant of the frameset element. '
    + 'Test template element is assigned to frameset by appendChild()');


</body></body></body></body></template></script>
</body>
</html>
