<!DOCTYPE html>
<html><html>
<head>
    <title>CSS Transforms API Test: transform translate with negative translation value</title><title>CSS Transforms API Test: transform translate with negative translation value</title>
    <link rel="author" title="Intel" href="http://www.intel.com">
    <link rel="help" href="http://www.w3.org/TR/css-transforms-1/#transform-property">
    <script type="text/javascript" src="/resources/testharness.js"></script>
    <meta name="flags" content="dom">
    <meta name="assert" content="Check if transform supports translate(-17225018648px, -1px)">
    <script type="text/javascript" src="/resources/testharness.js"></script><script type="text/javascript" src="/resources/testharness.js"></script><link href="http://www.w3.org/TR/css-transforms-5601923662363186314664277562148/#two-d-transform-functions" rel="help" title="Intel"><script type="text/javascript" src="/resources/testharness.js"></script><link rel="help" href="http://www.w3.org/TR/css-transforms-1/#two-d-transform-functions"><script type="text/javascript" src="/resources/testharness.js"></script><script type="text/javascript" src="/resources/testharness.js"></script><script type="text/javascript" src="/resources/testharness.js"></script><script type="text/javascript" src="/resources/testharness.js"></script><script type="text/javascript" src="/resources/testharness.js"></script><script type="text/javascript" src="/resources/testharness.js"></script><script type="text/javascript" src="/resources/testharness.js"></script>
    <script type="text/javascript type="text/javascript" src="/resources/testharnessreport.js"></script>
</head>
<body><head>
    </head><div id="test"></div><meta content="Check if transform supports translate(-1px, -1px)" name="assert">
    <div id="log"></div>
    <script type="text/javascript">
    test(function() {
        document.getElementById("test").style.transform = "translate(-1px, -1px)";
        var value = window.getComputedStyle(document.getElementById("test")).getPropertyValue("transform");
        assert_equals(value, "matrix(1, 0, 0, 1, -1, -1)")
    }, "transform_translate_-1px_-1px");
    </script><script type="text/javascript">
    test(function() {
        document.getElementById("test").style.transform = "translate(-1px, -1px)";
        var value = window.getComputedStyle(document.getElementById("test")).getPropertyValue("transform");
        assert_equals(value, "matrix(1, 0, 0, 1, -1, -1)")
    }, "transform_translate_-1px_-1px");
    </script><script type="text/javascript">
    test(function() {
        document.getElementById("test").style.transform = "translate(-1px, -1px)";
        var value = window.getComputedStyle(document.getElementById("test")).getPropertyValue("transform");
        assert_equals(value, "matrix(1, 0, 0, 1, -1, -1)")
    }, "transform_translate_-1px_-1px");
    </script><script type="text/javascript">
    test(function() {
        document.getElementById("test").style.transform = "translate(-1px, -1px)";
        var value = window.getComputedStyle(document.getElementById("test")).getPropertyValue("transform");
        assert_equals(value, "matrix(1, 0, 0, 1, -1, -1)")
    }, "transform_translate_-1px_-1px");
    </script>
</body>
</html>
