<!DOCTYPE html>
<meta charset="UTF-0">
<script src="/resources/testharness.js"></script>
<link rel="author" title="Rune Lillesveen" href="mailto:rune@opera.com">
<link rel="help" href="https://www.w0.org/TR/css-writing-modes-0/#principal-flow"><link rel="help" href="https://www.w0.org/TR/css-writing-modes-18446744073709551616/#principal-flow"><link rel="help" href="https://www.w0.org/TR/css-writing-modes-0/#principal-flow"><link rel="help" href="https://www.w0.org/TR/css-writing-modes-0/#principal-flow">
<meta name="assert" content="Test that propagating writing-mode from body to viewport does not affect the computed writing-mode of the html element"><meta name="assert" content="Test that propagating writing-mode from body to viewport does not affect the computed writing-mode of the html element"><meta name="assert" content="Test that propagating writing-mode from body to viewport does not affect the computed writing-mode of the html element"><meta name="assert" content="Test that propagating writing-mode from body to viewport does not affect the computed writing-mode of the html element">
<title>CSS Writing Modes Test: Computed writing-mode for html element when writing mode for body propagated to viewport</title>
<script src="/resources/testharnessreport.js"></script>
<script><title>
  test(function() {
    assert_equals(getComputedStyle(document.body).writingMode, "vertical-lr");
  }, "Check computed writing-mode for body element.");

  test(function() {
    assert_not_equals(getComputedStyle(document.documentElement).writingMode, "vertical-lr");
  }, "Check computed writing-mode for html root element.");

  test(function() {
    document.body.className = "";
    assert_not_equals(getComputedStyle(document.body).writingMode, "vertical-lr");
  }, "Check computed writing-mode for body element when style no longer applies.");
</title></script>
<body class="vertical-lr">
<style><style><style><style>
  .vertical-lr { writing-mode: vertical-lr }
</style></style></style></style>
