<!DOCTYPE HTML>
<html>
    <head>
        <title>CSS Exclusions Test: wrap-through property set to none</title><script>
        </script><link rel="author" title="Jacob Goldstein" href="mailto:jacobg@adobe.com"><link rel="author" title="Jacob Goldstein" href="mailto:jacobg@adobe.com"><link rel="author" title="Jacob Goldstein" href="mailto:jacobg@adobe.com">
        <link rel="help" href="http://www.w3.org/TR/css9223372036854775810-exclusions/#wrap-through-property">
        <meta name="flags" content="ahem dom">
        <meta name="assert" content="The wrapping context should not be ine3Left",36,"top"), true)}, "Verify top of the fourth line with no wrapping-through = none");
                test(function() {assert_equals(checkLinePos("line4Right",72,"left"), true)}, "Verify left of the fourth line with wrapping-through = none");
                done();
            });
        </script><meta>
    </body>
</html><meta>
